# Pandas

```python
from nose.tools import assert_true, assert_false, \
    assert_almost_equal, assert_equal, assert_raises
```

**Problem 1 (25 points).** Use Pandas to read in the file pruning_machine_learning_10292014.txt. Create a pandas DataFrame that has all the columns in the text file EXCEPT for "edge_label" and "minmaxdfe." Also all missing (NaN) values should be replaced with the mean value for the relevant column. **Hint**: read up on [``fillna``](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.fillna.html) and [``drop``](http://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.drop.html)

```python
import os
import pandas as pd
import matplotlib.pyplot as plt
from collections import defaultdict, OrderedDict
import datetime
import numpy as np
from dateutil import parser
```

<!-- grade_id: cell-3ba26c1abd0c4eac, points: 0, grade: false, solution: true -->
```python
#data1 = os.path.join("..","..","..","DATA","Numerics","mimic2","bp","subjects")
#os.path.exists(DATADIR)

#data1 = pd.read_table("data.txt")

#del data1["edge_label"]
#del data1["minmaxdfe"]
#data1.replace("None",np.NaN,inplace=True)
#data1.replace(np.NaN,data1.mean(),inplace=True)

#data1

data1 = pd.read_table("data.txt", na_values="None")
del data1["edge_label"]
del data1["minmaxdfe"]
data1.replace(np.NaN, data1.mean(), inplace=True)
data1

```

Output:
```
basename  prune_label  pathlength  exterior2surface  \
0     671_PV_skeleton_graphs.pckle            0          23          0.775000   
1     671_PV_skeleton_graphs.pckle            0          45          0.910923   
2     671_PV_skeleton_graphs.pckle            0         137          0.922793   
3     202_PV_skeleton_graphs.pckle            0           3          0.713208   
4     202_PV_skeleton_graphs.pckle            0           6          0.732123   
5     202_PV_skeleton_graphs.pckle            0         113          0.922076   
6     849_PV_skeleton_graphs.pckle            0          21          0.658359   
7     849_PV_skeleton_graphs.pckle            0          17          0.742336   
8     849_PV_skeleton_graphs.pckle            0          22          0.847027   
9     849_PV_skeleton_graphs.pckle            0          18          0.622440   
10    849_PV_skeleton_graphs.pckle            0          14          0.624347   
11    849_PV_skeleton_graphs.pckle            0       
```

<!-- grade_id: cell-092a4c8cba81de8e, points: 5, grade: true, solution: false -->
```python
assert_equal(type(data1), pd.core.frame.DataFrame)
```

<!-- grade_id: cell-4025eedd103f6b1c, points: 5, grade: true, solution: false -->
```python
assert_equal(data1.shape, (753, 10))
```

<!-- grade_id: cell-fc880e6d432d8859, points: 0, grade: true, solution: false -->
```python
assert_false(pd.isnull(data1["exterior2surface"][70]))
```

<!-- grade_id: cell-b60e4a1a17c357b4, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-5ec8fea30bec179b, points: 3, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-fe95326b65ea2b0a, points: 0, grade: true, solution: false -->
```python
assert_almost_equal(data1["mindfe"][732], 4.1495063469675602)
```

<!-- grade_id: cell-c45f0ad5ec2adcf6, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-b4c46ca32e766387, points: 0, grade: true, solution: false -->
```python
assert_true("maxdfe" in data1.columns)
```

<!-- grade_id: cell-35be7e5c8a2efdd7, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-fbc94f76dff38ab4, points: 0, grade: true, solution: false -->
```python
assert_false("edge_label" in data1.columns)
```

<!-- grade_id: cell-754d3dba1cdcf3a2, points: 3, grade: true, solution: false -->
```python

```

**Problem 2 (10 points):** Write a function `get_date` that takes a single positional argument and uses the parser defined in the [dateutil](http://dateutil.readthedocs.io/en/stable/parser.html#dateutil.parser.parse) package to convert the argument to a [`datetime`](https://docs.python.org/3/library/datetime.html) [`date`](https://docs.python.org/3/library/datetime.html#datetime.date) object. If the argument cannot be converted to a `date` object, the function should return a [Pandas `NaT`](https://pandas.pydata.org/pandas-docs/stable/missing_data.html#datetimes) (not a time) object.


<!-- grade_id: cell-5db86753d6569855, points: 0, grade: false, solution: true -->
```python
def get_date(t):
    #t= set(t)
    numbers = {'0','1','2','3','4','5','6','7','8','9','/'}
    if set(t).issubset(numbers):
        
        return parser.parse(t).date()
    else:
        return np.datetime64('NaT')
```

<!-- grade_id: cell-a8491be3cd38b9d5, points: 0, grade: true, solution: false -->
```python
assert_equal(get_date("5/17/2007"), datetime.date(2007, 5, 17))
```

<!-- grade_id: cell-6bb0dc2a182992bb, points: 0, grade: true, solution: false -->
```python
assert_true(pd.isnull(get_date("Hello")))
```

<!-- grade_id: cell-304f2dd43a0f9976, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-95ecb610a1de2481, points: 5, grade: true, solution: false -->
```python

```

**Problem 3 (5 points):** Write a function get_date that takes a single positional argument and uses the parser defined in the dateutil package to convert the argument to a datetime time object. If the argument cannot be converted to a time object, the function should return a Pandas NaT (not a time) object.

<!-- grade_id: cell-ed5f1125c841cb25, points: 0, grade: false, solution: true -->
```python
def get_time (ti):
    #t= set(t)
    numbers = {'0','1','2','3','4','5','6','7','8','9',':','A','M','P'}
    if set(ti).issubset(numbers):
        
        return parser.parse(ti).time()
    else:
        return np.datetime64('NaT')
```

<!-- grade_id: cell-81bf30ac02a094bf, points: 2, grade: true, solution: false -->
```python
assert_true(pd.isnull(get_time("Hello")))
```

<!-- grade_id: cell-1c814383db8fce36, points: 3, grade: true, solution: false -->
```python

```

**Problem 4 (15 points):** Write a function `get_range` that takes a positional argument containing a string from which to extract a range and a keyword argument named `delimiter` that has a string indicating the delimiter between the low and high values of the range. The function should **try** to return a tuple of two floating point values (the low range value and the high range value). 

The function should return a two-tuple `np.nan` values  (`(np.nan, np.nan)`) for the following conditions:

1. The string does not contain exactly one instance of `delimiter`.
1. The attempt to extract two floating point values fails.   

If either the positional or keyword arguments are not strings, return a `TypeError`.

<!-- grade_id: cell-e8ca21293ad89c6d, points: 0, grade: false, solution: true -->
```python
import numbers

def get_range(W,delimiter="-"):
    if type(W) != str or type(delimiter) != str:
        raise TypeError("arguments need to be strings")

    
    W = W.split(delimiter)
    if len(W) != 2:
        return (np.nan, np.nan)
    try:
        return (float(W[0]), float(W[-1]))
    except:
        return (np.nan, np.nan)
```

<!-- grade_id: cell-3e45e066448be1f0, points: 2, grade: true, solution: false -->
```python
assert_raises(TypeError, get_range, 5.4)
```

<!-- grade_id: cell-74138019787f16f8, points: 3, grade: true, solution: false -->
```python
assert_raises(TypeError, get_range, 5.4, delimiter=(5,4,3))
```

<!-- grade_id: cell-7169e27bcd70d566, points: 0, grade: true, solution: false -->
```python
assert_true(np.isnan(get_range("12-24-1972")).all())
```

<!-- grade_id: cell-593d149049688ba0, points: 0, grade: true, solution: false -->
```python
assert_equal(get_range("15-20"), (15,20))
```

<!-- grade_id: cell-c563f84b6669df5d, points: 3, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bd3a94e5c9ca5805, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0b550ff356f19caa, points: 5, grade: true, solution: false -->
```python

```

**Problem 5 (10 points):** Write a function `is_number` that takes a single positional argument and returns True if that argument can be converted to a `float` and `false` if it cannot.

<!-- grade_id: cell-e5775eac7f9cf462, points: 0, grade: false, solution: true -->
```python
def is_number(cutup):
    try:
        float(cutup)
        return True
    except ValueError:
        return False
```

<!-- grade_id: cell-5a5e78ec1a999d97, points: 0, grade: true, solution: false -->
```python
assert_true(is_number("5.3e4"))
```

<!-- grade_id: cell-ece8aebfca0bfaeb, points: 0, grade: true, solution: false -->
```python
assert_false(is_number(">5000"))
```

<!-- grade_id: cell-5fd7f943c02e72bd, points: 5, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-567f714caf748bb3, points: 5, grade: true, solution: false -->
```python

```

**Problem 6 (5 points):** Write a function `get_value` that takes a single positional argument (e.g. `x`) and returns `x` converted to a float. If the conversion fails, return `np.nan`.

<!-- grade_id: cell-2dffeb667ae8c468, points: 0, grade: false, solution: true -->
```python
def get_value(value):
    try:
        float(value)
        return float(value)
    except ValueError:
        return np.nan

```

<!-- grade_id: cell-054827c5de345630, points: 2, grade: true, solution: false -->
```python
assert_true(np.isnan(get_value("[5.4]")))
```

<!-- grade_id: cell-e4e569a39d5f20aa, points: 3, grade: true, solution: false -->
```python
assert_equal(get_value("5.7"), 5.7)
```

### For the following problems we will use the `sampleLabs.txt` file
#### In this directory is a file (`sampleLabs.txt`) with a set of lab values. The file includes lab values obtained on different individuals during their in-patient hospitalizations. 

**Problem 7 (15 points)**: Use `is_number()`, list comprehension, and the Pandas Series `unique` method to create a list of all non-numeric values in the `result` Series of `lab_data`. By non-numeric, I mean any value that cannot be converted to a `float`.

<!-- grade_id: cell-3456eb0295551017, points: 0, grade: false, solution: true -->
```python
#non_number_results = None
#lab_data = pd.read_table("samplelabs.txt")
#LD = lab_data.values.tolist()

#non_number_results = [j for i in LD for j in pd.unique(i) if is_number(j)==False]


non_number_results = None
lab_data = pd.read_table("samplelabs.txt").result.unique()
non_number_results = [i for i in lab_data if not is_number(i)]

```

<!-- grade_id: cell-556d11607452a963, points: 5, grade: true, solution: false -->
```python
assert_true(type(non_number_results)==list)
```

<!-- grade_id: cell-161eccac12e719b5, points: 10, grade: true, solution: false -->
```python

```

**Problem 8. (40 Points):**  Use Pandas to write a function named `read_labs` that does the following:

* Takes the name of the file to read as a positional argument
* Takes as a keyword argument named `converters` a dictionary with keys equal to the test component names (columns in DataFrame) and as values, functions that convert the expected column value to the appropriate type. You should use the ordered dictionary (`converters`) I define below.
* Reads the contents of the input file into a **default dictionary** keyed by **visitid**. The value for each **visitid** should be a **list** of tests. For each row, use `converters` to convert the row value from a string to the appropriate data type (eg. float, datetime.date). Each test result result should be stored as a tuple of 2-tuples with the first element being the explanatory metadata (i.e, the column name) and the second element being the converted value.

#### Hints

1. Use the DataFrame `iterrows` method to iterate over each row in the DataFrame.
1. Use the function `get_test_values` to convert the data types of each row.
1. Use slicing on the row to grab the test values (`visitid` is going to be the key for the dictionary and should not be part of the test result tuple.
1. A row is a Pandas Series and we can access the elements of the row using either an index or the column name:

```Python
print(row)
print()
print(row[2])
print(row["ctime"])

visitid    OMHioJh8XEeq7152
cdate             6/13/2007
ctime                 06:30
pqno       1181750718122403
test                  CREAT
result                  1.0
unit                  mg/dL
range               0.5-1.4
Name: 5, dtype: object

06:30
06:30
```

Your resulting output should look something like this:

```Python
defaultdict(list,
            {'+yhZLyY5Uqra5115': [[('cdate', datetime.date(2005, 4, 29)),
               ('ctime', datetime.time(6, 30)),
               ('pqno', 1114780124780442),
               ('test', 'CREAT'),
               ('result', 3.4),
               ('unit', 'mg/dL'),
               ('range', (0.5, 1.4))],
              [('cdate', datetime.date(2005, 4, 28)),
               ('ctime', datetime.time(6, 45)),
               ('pqno', 1114692308737109),
               ('test', 'CREAT'),
               ('result', 8.6),
               ('unit', 'mg/dL'),
               ('range', (0.5, 1.4))],
              [('cdate', datetime.date(2005, 5, 2)),
               ('ctime', datetime.time(4, 0)),
               ('pqno', 1115041801503288),
               ('test', 'CREAT'),
               ('result', 1.6),
               ('unit', 'mg/dL'),
               ('range', (0.5, 1.4))],
              [('cdate', datetime.date(2005, 4, 30)),
               ('ctime', datetime.time(4, 0)),
               ('pqno', 1114875733484671),
               ('test', 'CREAT'),
               ('result', 1.9),
               ('unit', 'mg/dL'),
               ('range', (0.5, 1.4))],
              [('cdate', datetime.date(2005, 4, 26)),
               ('ctime', datetime.time(16, 5)),
               ('pqno', 1114551787814371),
               ('test', 'CREAT'),
               ('result', 16.0),
               ('unit', 'mg/dL'),
               ('range', (0.5, 1.4))],
              [('cdate', datetime.date(2005, 5, 1)),
               ('ctime', datetime.time(4, 0)),
               ('pqno', 1114957388877371),
               ('test', 'CREAT'),
               ('result', 1.7),
               ('unit', 'mg/dL'),
               ('range', (0.5, 1.4))]],}
```

```python
from collections import OrderedDict
converters = OrderedDict((('cdate', get_date), 
                          ('ctime', get_time), 
                          ('pqno', lambda x:x), 
                          ('test', lambda x:x), 
                          ('result', get_value), 
                          ('unit', lambda x:x), 
                          ('range', get_range)))
```

<!-- grade_id: cell-03d7de4141bc1410, points: 0, grade: false, solution: true -->
```python
import datetime
import os
import pandas as pd
from collections import defaultdict

        
def get_test_values(row, converters):
    """
    Arguments:
        row: A row from a Pandas DataFrame
        converters: A Dictionary with keys the DataFrame column names and 
            values functions to convert each column value
    Returns:
        A tuple of two-tuples. Each two-tuple consists of the column name
            and the converted column value
    """
    return [(c,v(row[c])) for c,v in converters.items()]
    
    
def read_labs(fname, converters):
    """
    Arguments:
        fname: Path to a file containing lab values
        converters: A Dictionary with keys the DataFrame column names and 
            values functions to convert each column value
    Returns:
        A default dictionary with keys being the visitid and values
            begin a list of test results
    """
    FORTY = defaultdict(list)
    df = pd.read_table(fname)
   
    for index, r in df.iterrows():
        
                key = r[0]
                value=get_test_values(r[1:], converters)
                
                
                FORTY[key].append(value) 
                
                
    return FORTY
    
    
```

```python
sorted_labs = read_labs("samplelabs.txt", converters)
```

<!-- grade_id: cell-8ced45219fdd9711, points: 5, grade: true, solution: false -->
```python
assert_equal(sorted_labs['X4HZ5/mbdrls7055'][8][1][1], datetime.time(4, 4))
```

<!-- grade_id: cell-54f1522d387b02f1, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(sorted_labs["3WobBYVBHdIa6106"][5][6][1],(0.5,1.4))
```

<!-- grade_id: cell-3aa0b756312408ff, points: 5, grade: true, solution: false -->
```python
assert_true('QLE6UIMptiKf5228' in sorted_labs)
```

<!-- grade_id: cell-996eb2c6d4ce08ea, points: 5, grade: true, solution: false -->
```python
assert_equal(len(sorted_labs['n9b8R+Aynzy+5257']),  31)
```

<!-- grade_id: cell-b23ec7c9eb0d9b9a, points: 5, grade: true, solution: false -->
```python
assert_equal(sorted_labs['yyutduwkJpYJ7131'][1][1], 
             ('ctime', datetime.time(6, 15)))
```

<!-- grade_id: cell-81c651eb76ec7312, points: 5, grade: true, solution: false -->
```python
assert_true(type(sorted_labs['yyutduwkJpYJ7131'][1]),tuple)
```

<!-- grade_id: cell-8ebc48c309376589, points: 5, grade: true, solution: false -->
```python
assert_true(np.isnan(sorted_labs['deaOcpSo0N617039'][14][4][1]))
```

<!-- grade_id: cell-4e60f6362afad4f4, points: 5, grade: true, solution: false -->
```python
assert_true(np.isnan(sorted_labs['QLE6UIMptiKf5228'][17][6][1]).all())
```

**Problem 9 (20 points)**: Use Pandas to read in the file ~/Data/Numerics/mimic2/bp/subjects/19468.txt. Rename the columns to be "systolic" and "diastolic". Create a new column named "difference" with values equal to the difference between the "systolic" and "diastolic" values.

<!-- grade_id: cell-f45413b9622cd5dc, points: 0, grade: false, solution: true -->
```python
import os
import sqlite3 as sqlite
DATADIR = os.path.join(os.path.expanduser("~"),"DATA","Numerics","mimic2","bp","subjects")
os.path.exists(DATADIR)

bpdata = pd.read_table(os.path.join(DATADIR,"19468.txt"),header=None)

bpdata.columns=["systolic","diastolic"]

bpdata["difference"] = bpdata["systolic"] - bpdata["diastolic"]

bpdata
```

Output:
```
systolic  diastolic  difference
0        179         75         104
1        169         69         100
2        185         76         109
3        176         76         100
4        130         76          54
5        188         74         114
6        138         62          76
7        123         55          68
8        125         56          69
9        112         51          61
10       106         48          58
11       109         60          49
12       113         57          56
13       121         56          65
14       124         60          64
15       133         62          71
16       124         79          45
17       156         69          87
18       122         67          55
19        93         75          18
20       165         72          93
21       155         73          82
22       117         61          56
23       129         70          59
24       168         71          97
25       147         72          75
26       133         65         
```

<!-- grade_id: cell-34db8c293edf535b, points: 5, grade: true, solution: false -->
```python
assert_equal(bpdata.shape, (35,3))
```

<!-- grade_id: cell-a3b0e7ebb62f2ffa, points: 5, grade: true, solution: false -->
```python
assert_equal(bpdata.columns.tolist(), ['systolic', 'diastolic', 'difference'])
```

<!-- grade_id: cell-829c6b2b475fda3f, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(bpdata["systolic"].mean(), 142.71428571428572)
```

<!-- grade_id: cell-1777b4334d4cb9d2, points: 5, grade: true, solution: false -->
```python
assert_equal(bpdata["systolic"].dtype, np.int64)
```