# M4-Collections (Part 1)

```python
from nose.tools import assert_equal
```

**Problems 1 and 2** use the variable ``report`` defined below.

```python
report = """DATE: [**2644-1-17**] 10:43 AM
     CT HEAD W/O CONTRAST                                            Clip # [**Clip Number (Radiology) 12567**]
     Reason: ? stability of bleed                                        
     Admitting Diagnosis: HEAD BLEED
     ______________________________________________________________________________
     UNDERLYING MEDICAL CONDITION:
      [**Age over 90 **] year old woman with ? hemorragic mass                                        
     REASON FOR THIS EXAMINATION:
      ? stability of bleed                                                            
     No contraindications for IV contrast
     ______________________________________________________________________________
                                     FINAL REPORT  (REVISED)
     HISTORY:  [**Age over 90 **] year-old woman with a history of metastatic lung cancer to the
     brain.  Question hemorrhagic mass.  Question stability of bleed.
     
     COMPARISON:  Head CT from [**Hospital 12568**] Hospital dated [**2644-1-16**].
     
     TECHNIQUE:  Noncontrast head CT.
     
     FINDINGS:  Compared to the prior days study, there is stable appearance of the
     right parietal intraparenchymal hemorrhage with surrounding edema.  At the
     superior margin of the parenchymal hemorrhage there is a rounded heterogeneous
     focus which could represent a metastatic lesion.  An additional 2mm hyperdense
     focus, possibly hemorrhage, is noted in the posteromedial margin of the left
     thalamus, with surroundng edema.  Low-attenuation foci seen in both basal
     ganglia and insular regions are consistent with chronic lacunar infarcts.
     There is no shift of midline structures.  The ventricles are stable in
     appearance.  The osseous and soft tissue structures are unremarkable.
     
     IMPRESSION:  Stable appearance of right parietal lobe and left thalamic
     hemorrhages, which are concerning for hemorrhagic metastasis in this patient
     with known metastatic lung carcinoma to the brain."""
```

**Problem 1 (20 points):** For variable report above, use the string ``find()`` (or ``index()``) method and slicing to modify `get_history` and `get_impression` to return substrings strings that contain (1) the history section (i.e. the text from "HISTORY:" (inclusive) to "TECHNIQUE:" (exclusive), and (2) the Impression section (the text from "IMPRESSION:" to the end of the report) respectively.

<!-- grade_id: cell-7a64746c185a6f27, points: 0, grade: false, solution: true -->
```python
def get_history(report):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-a3c3e6f0a4ad87e5, points: 0, grade: true, solution: false -->
```python
assert_equal(get_history(report), 'HISTORY:  [**Age over 90 **] year-old woman with a history of metastatic lung cancer to the\n     brain.  Question hemorrhagic mass.  Question stability of bleed.\n     \n     COMPARISON:  Head CT from [**Hospital 12568**] Hospital dated [**2644-1-16**].\n     \n     ')
```

<!-- grade_id: cell-e947cd9ab169fb66, points: 0, grade: false, solution: true -->
```python
def get_impression(report):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-0d9c212551a833be, points: 0, grade: true, solution: false -->
```python
assert_equal(get_impression(report), 'IMPRESSION:  Stable appearance of right parietal lobe and left thalamic\n     hemorrhages, which are concerning for hemorrhagic metastasis in this patient\n     with known metastatic lung carcinoma to the brain.')
```

<!-- grade_id: cell-4ebf5e3b2e90c01f, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-39ddf3158619f8c9, points: 10, grade: true, solution: false -->
```python

```

**Problem 2 (20 points):** Ignoring case and periods (e.g. "``present``" and "``present.``" should not be counted as two distinct words), how many unique words are in the report? In the sense of this problem a "word" is a collection of characters separated by spaces from other words. So "`[**3352-11-25**]`" is a word.

<!-- grade_id: cell-2663b43251c92d07, points: 0, grade: false, solution: true -->
```python
num_words = -1
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-81b91548b519a29c, points: 20, grade: true, solution: false -->
```python

```

**Problem 3 (20 points):** In the table below are actual values from a metabolic lab test. Define a named tuple named ``labresult`` for a lab result that can store an individual test result where each test value consists of a "component", "value", and "standard range". Create an instance of ``labresult`` for each row in the table below. Put all the instances of ``labresult`` in an appropriate collection (``list``, ``tuple``, or ``dictionary``). Explain your choice for the type of collection.

|Component          |      Value |   Standard Range | 
|-------------------|---------------|---------------- | 
|Sodium            |     139 mmol/L  |  136 - 144 mmol/L |
|Potassium, S/P    |     4.3 mmol/L  |  3.3 - 5.0 mmol/L |
|Chloride, S/P     |     107 mmol/L  |  102 - 110 mmol/L |
|Carbon Dioxide, S/P |   22 mmol/L   |  20 - 26 mmol/L |
|BUN, S/P           |    16 mg/dL    |  8 - 24 mg/dL |

<!-- grade_id: cell-15f1090f4f292267, points: 20, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Problem 4 (15 points).** Given the following lists
```Python
list1 = [1, 2, 3]
list2 = [3, 4, 5]
list3 = [6, 7, 8]
```
use slicing, Python builtin functions, and list methods to create the following lists using a single Python statement for each problem:

1. ``list4 = [1, 2, 3, 4, 5, 6, 7]``
1. ``list5 = [2, 4, 6, 8]``
1. ``list6 = [8, 7, 6, [3, 4, 5]]``

<!-- grade_id: cell-20ffce6459902f4a, points: 15, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Problem 5 (10 points):** Create a Python set that satisfies the following definition:

$$S = \{x: x \% 3=0 \text{ and} -10 \le x \le 10 \text{ and } x \in \mathbb{Z}\}$$



<!-- grade_id: cell-a332a46a0c99ad21, points: 0, grade: false, solution: true -->
```python
set5 = {}
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5d2a7025db9004db, points: 10, grade: true, solution: false -->
```python

```

**Problem 6 (10 points): Create a Python set that satisfies the following definition:**

*The set of all pairs of integers between 0 and 3 inclusive where the order of the pairs matters*

<!-- grade_id: cell-e0dad53fcc8a59f4, points: 0, grade: false, solution: true -->
```python
set_problem6 = {}
# YOUR CODE HERE
raise NotImplementedError()
type(set_problem6)
```

<!-- grade_id: cell-4662631670c7df00, points: 10, grade: true, solution: false -->
```python

```

**Problem 6 (10 points):** Given the list ``list1`` 
```Python
list1 = ["a", (2,3), 4, []]
```
write a single Python expression that results in a new variable x=4 and modifies ``list1`` to be
```Python
list1 = ["a", (2,3), []]
```

<!-- grade_id: cell-6c9b9da6196c9f3b, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Problem 7 (20 points):** In the context of your actual or potential term project describe ways that you could potentially use 

* Python lists
* Python tuples (or named tuples)
* Python sets

<!-- grade_id: cell-3816da0112843826, points: 20, grade: true, solution: true -->
YOUR ANSWER HERE