# Dictionaries

This homework has you work with Python dictionaries and a couple of variants, namely default dictionaries and counters, both defined in the Python library [collections](https://docs.python.org/3/library/collections.html). These problems will require you to use if/else statments and for loops which I assume you are familiar with from your CodeAcademy review. In case you want more practice with these, I've put notebooks in the Resources directory. The problems get progressively harder as you move towards more and more primative/basic collections.

**Problems 1 and 2** use the variable ``report`` defined below.

```python
from nose.tools import assert_equal, assert_true, assert_false
```

```python
report = """
Patient: DOE, JOHN
Referring Physician: DR. DAVID LIVESEY
MRN : JD4USARAD DOB: 01/01/1961
Exam Date: 06/05/2010
FAX: (305) 418-8166
    
CLINICAL HISTORY: ABDOMINAL PAIN.
    
TECHNIQUE: Multiple axial CT images were obtained through the abdomen and pelvis after
administration of oral contrast material only.

COMMENTS:
There is evidence of diffuse hepatic hypoattenuation compatible with fatty infiltration. There is no
intra or extrahepatic biliary ductal dilatation. The patient is status post cholecystectomy. The spleen
is normal. The pancreas is of normal contour and attenuation characteristics. There is no evidence of
adrenal mass.

Moderate sized fat containing supraumbilical hernia is present.
The kidneys are normal in size, shape and configuration. No renal or ureteral calculi are identified.
There is no hydroureter or hydronephrosis.
There is no evidence for appendicitis. Several fluid-filled loops of small bowel are present compatible
with mild enteritis. There is no bowel wall thickening. No evidence for small or large bowel
obstruction. There is no evidence of abdominal ascites or lymphadenopathy.
There is no evidence of intrinsic or extrinsic bladder mass. There is no pelvic ascites or
lymphadenopathy.
The uterus and ovaries are grossly unremarkable.
Images of the lung bases show no evidence of pleural or parenchymal mass. There are no pleural
effusions. Scarring is present in the right middle lobe and lingula as well as both lung bases.
The bony structures are free of lytic or blastic lesions. Multilevel degenerative changes are seen
involving the thoracolumbar spine.
Scattered calcifications are seen involving the aorta and major branches compatible with
atherosclerosis.

IMPRESSION:
ABDOMEN:
1. Several fluid-filled loops of small bowel are present compatible with mild enteritis.
2. Fatty liver.
3. Status post cholecystectomy.
4. Fat containing umbilical hernia.
PELVIS:
1. No evidence of diverticulitis or acute inflammatory process in the pelvis.
Discussed with Dr. DAVID LIVESEY
-Electronically Signed by: RADIOLOGIST, ADMIN on 06/05/2010 2:31:33 PM"""
```

**Problem 1 (20 points):** Use the ``Counter`` class defined in ``collections`` to modify the function `def count_word_v1` to count the frequency of words in ``report``. Before counting the number of words, convert everything to uppercase and replace periods (".") and commas (",") with empty spaces. The function takes as an argument a string containing a report and returns a Counter object.

<!-- grade_id: cell-ca5c0aca364ed301, points: 0, grade: false, solution: true -->
```python
from collections import Counter

def count_words_v1(x):
    wordcount = Counter(x.upper().replace(". "," ").replace('."'," ").replace(".'", " ").replace(": "," ").replace(":\n"," ").replace(", "," ").replace(".\n"," ").replace("\n"," ").replace("\t"," ").replace("**"," ").replace(".","").split())
    return wordcount

count_words_v1(report)


```

Output:
```
Counter({'(305)': 1,
         '-ELECTRONICALLY': 1,
         '01/01/1961': 1,
         '06/05/2010': 2,
         '1': 2,
         '2': 1,
         '2:31:33': 1,
         '3': 1,
         '4': 1,
         '418-8166': 1,
         'ABDOMEN': 2,
         'ABDOMINAL': 2,
         'ACUTE': 1,
         'ADMIN': 1,
         'ADMINISTRATION': 1,
         'ADRENAL': 1,
         'AFTER': 1,
         'AND': 6,
         'AORTA': 1,
         'APPENDICITIS': 1,
         'ARE': 9,
         'AS': 2,
         'ASCITES': 2,
         'ATHEROSCLEROSIS': 1,
         'ATTENUATION': 1,
         'AXIAL': 1,
         'BASES': 2,
         'BILIARY': 1,
         'BLADDER': 1,
         'BLASTIC': 1,
         'BONY': 1,
         'BOTH': 1,
         'BOWEL': 4,
         'BRANCHES': 1,
         'BY': 1,
         'CALCIFICATIONS': 1,
         'CALCULI': 1,
         'CHANGES': 1,
         'CHARACTERISTICS': 1,
         'CHOLECYSTECTOMY': 2,
         'CLINICAL': 1,
         'COMMENTS': 1,
         'COMPATIBLE': 4,
     
```

<!-- grade_id: cell-9e4147133f20571a, points: 15, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-5472a098a8b1a265, points: 5, grade: true, solution: false -->
```python
assert_true(type(count_words_v1(report))==Counter)

```

**Problem 2 (30 points):** Repeat problem one but instead of using a `Counter`, use the ``defaultdict`` defined in the ``colletions`` library.

<!-- grade_id: cell-15995fa74f7d5ce5, points: 0, grade: false, solution: true -->
```python
from collections import defaultdict

def count_words_v2(x):
    x = x.upper().replace(". "," ").replace('."'," ").replace(".'", " ").replace(": "," ").replace(":\n"," ").replace(", "," ").replace(".\n"," ").replace("\n"," ").replace("\t"," ").replace("**"," ").replace(".","").split()    
    WC = defaultdict(lambda:0)
    wordset = set(x)
    for word in wordset:
        WC[word]+=1
        
    return WC

count_words_v2(report)

```

Output:
```
defaultdict(<function __main__.count_words_v2.<locals>.<lambda>>,
            {'(305)': 1,
             '-ELECTRONICALLY': 1,
             '01/01/1961': 1,
             '06/05/2010': 1,
             '1': 1,
             '2': 1,
             '2:31:33': 1,
             '3': 1,
             '4': 1,
             '418-8166': 1,
             'ABDOMEN': 1,
             'ABDOMINAL': 1,
             'ACUTE': 1,
             'ADMIN': 1,
             'ADMINISTRATION': 1,
             'ADRENAL': 1,
             'AFTER': 1,
             'AND': 1,
             'AORTA': 1,
             'APPENDICITIS': 1,
             'ARE': 1,
             'AS': 1,
             'ASCITES': 1,
             'ATHEROSCLEROSIS': 1,
             'ATTENUATION': 1,
             'AXIAL': 1,
             'BASES': 1,
             'BILIARY': 1,
             'BLADDER': 1,
             'BLASTIC': 1,
             'BONY': 1,
             'BOTH': 1,
             'BOWEL': 1,
             'BRANCHES': 1,
             'BY': 1,
           
```

<!-- grade_id: cell-e56cdb614c78eb19, points: 25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-1edbfbd3d33241c2, points: 5, grade: true, solution: false -->
```python
assert_true(type(count_words_v2(report))==defaultdict)

```

**Problem 3 (50 points):** Modify `count_kmers` below as follows:

### Arguments

* `count_kmers` takes as a **positional argument** named `sequences` a dictionary of sequences
* `count_kmers` takes as a **keyword_argument** named `kmers` a tuple of kmers. The default tuple should be a tuple with the single kmer 'A'

### Return value

* `count_kmers` returns a dictionary with the keys corresponding to the keys of `sequences` and as values a dictionary with keys the valid kmers (valid DNA kmer) provided in the keyword argument `kmers` and as values the frequency count for that kmer in the corresponding sequence. (We will be returning a dictionary of dictionaries.

```python
{"seq1":{"CATG":count}}
{"seq2":{"CATG":count}}
{"seq3":{"CATG":count}}
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-15-e363061ff2b6> in <module>()
----> 1 {"seq1":{"CATG":count}}
      2 {"seq2":{"CATG":count}}
      3 {"seq3":{"CATG":count}}

NameError: name 'count' is not defined
```

```python
# the "\" character can be used as a line continuation in PPython
sequence1 = \
"""AGTTGTTAGTCTACGTGGACCGACAAGAACAGTTTCGAATCGGAAGCTTGCTTAACGTAGTTCTAACAGT
TTTTTATTAGAGAGCAGATCTCTGATGAACAACCAACGGAAAAAGACGGGTCGACCGTCTTTCAATATGC"""
sequence2= \
"""TGAAACGCGCGAGAAACCGCGTGTCAACTGTTTCACAGTTGGCGAAGAGATTCTCAAAAGGATTGCTTTC
AGGCCAAGGACCCATGAAATTGGTGATGGCTTTTATAGCATTCCTAAGATTTCTAGCCATACC"""
sequence3= \
"""GCAGGAATTTTGGCTAGATGGGGCTCATTCAAGAAGAATGGAGCGATCAAAGTGTTACGGGTCAAGA
AAGAAATCTCAAACATGTTGAACATAATGAACAGGAGGAAAAGATCTGTGACCATGCTCCTCATGCTGCT"""
sequence4= \
"""GCCCACAGCCCTGGCGTTCCATCTGACCACCCGAGGGGGAGAGCCGCACATGATAGTTAGCAAGCAGGAA
AGAGGAAAATCACTTTTGTTTAAGACCTCTGCAGGTGTCAACATGTGCACCCTTATTGCAATGGATTTGG
GAGAGTTATGTGAGGACACAATGACCTACAAATGCCCCCGGATCACTGAGACGGAACCAGATGACGTTGA
CTGTTGGTGCAATGCCACGGAGACATGGGTGACCTATGGAACATGTTCTCAAACTGGTGAACACCGACGA
GACAAACGTTCCGTCGCACTGGCACCACACGTAGGGCTTGGTCTAGAAACAAGAACCGAAACGTGGATGT"""

```

```python
sequences = {"seq1":sequence1, "seq2":sequence2, "seq3":sequence3}
def count_kmers(sequences,kmers=("A",)):
    
    DNA_set = {"A","T","G","C"}
    dict1 = {}
    
        #kmers = set(kmers)
        
    for s in sequences:
        #seq=seq.remove("/n")
        current_sequence = sequences[s] 
        dict2={}
        for kmer in set(kmers):
            if set(kmer).issubset(DNA_set):

                a = current_sequence.count(kmer)

                dict2[kmer] = a

        #values[a] += 1

        dict1[s] = dict2
    
    return dict1
    

count_kmers(sequences,kmers=("ACTG", "CAAC", "CBG", "GGG"))
```

Output:
```
{'seq1': {'ACTG': 0, 'CAAC': 2, 'GGG': 1},
 'seq2': {'ACTG': 1, 'CAAC': 1, 'GGG': 0},
 'seq3': {'ACTG': 0, 'CAAC': 0, 'GGG': 2}}
```

<!-- grade_id: cell-9e502a6f6df8ba6d, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-26e41e1487f1c9ea, points: 0, grade: true, solution: false -->
```python
sequences = {"seq1":sequence1, "seq2":sequence2, "seq3":sequence3}
rslts = count_kmers(sequences,kmers=("ACTG", "CAAC", "CBG", "GGG"))
assert_true(type(rslts)==dict)
assert_true("seq1" in rslts)
assert_false("CGB" in rslts["seq1"])
assert_equal(rslts["seq1"]["GGG"],1)
```

<!-- grade_id: cell-e63bfcabe80c7acc, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-d0aa3fc870398bf1, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-84bd1f44222f31ca, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ff0772453948090b, points: 10, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-910e692147050370, points: 10, grade: true, solution: false -->
```python

```

```python

```