# Differential Calculus


```python
from sympy import symbols
from sympy.plotting import plot
from sympy import diff
from sympy import sqrt, exp, cos, sin, pi, ln
from sympy import lambdify
import numpy as np
import numbers
import matplotlib.pyplot as plt
from nose.tools import assert_almost_equal, assert_true, assert_equal, assert_raises
from numpy.testing import assert_allclose
```

```python
x = symbols('x')
def func0(x):
    return x**3
def func1(x):
    """returns a third-order polynomial in x"""
    return sin(2*pi*x/6)*(x-3)**3-(5*(x+2)**2)+4*x*exp(-x**2)

def func2(x):
    """returns a third-order polynomial in x"""
    return (x-3)**3-5*(x+2)**2+4*x
```

```python
p = plot(func1(x), diff(func1(x)), (x, -10, 10), show=False, legend=True)
p[0].line_color = 'red'
p[0].label = "f(x)"
p[1].line_color = 'blue'
p[1].label="f'(x)"
p.show()
```

## Sympy [``lambdify``](http://docs.sympy.org/latest/modules/utilities/lambdify.html)

Sympy has a function ``lambdify`` that we can use to generate Python functions from Sympy expressions. We will create Python function versions of ``func1`` and its first derivative that we can use with numpy and matplotlib.

```python
fp1 = lambdify(x, func1(x), "numpy")
dfp1 = lambdify(x,diff(func1(x)), "numpy")
tt = np.arange(-10, 10, 0.1)
plt.plot(tt,fp1(tt))
```

**Problem 1. (25 points):** Write a function named ``d1a`` that computes the first derivative of a 1-D numpy array using a backward difference ($f_i'\approx \frac{f_i-f_{i-1}}{h}$) approximation of the derivative. The function should use numpy array slicing rather than a ``for`` loop. The resulting signal should have the same length as the input with a zero for the value where you did not compute the derivative. The function should take a numpy array as a positional argument and have a keyword argument ``h=1`` that provides the unit step size (i.e. what the length of one step in the array represents. 

Python has an [``isinstance``](https://docs.python.org/3/library/functions.html#isinstance) function that returns True if a variable is an instance of a type (e.g. ``str``) and False if it is not. The Python [``numbers``](https://docs.python.org/3.5/library/numbers.html) module provides definitions of number types, including a generic ``Number`` type of which all numbers are an instance. Use ``isinstance`` in conjunction with the ``numpy`` package and the ``numbers`` module to check if ``s`` is a numpy array and that ``h`` is a number. For both tests raise a ``TypeError`` if the user did not provide the corrected data type. Raise a `TypeError` if `h` is not positive. The function should have a docstring. Test the function with ``fp1``.

**Hints:** 

1. Use the [``numpy.zeros``](http://docs.scipy.org/doc/numpy/reference/generated/numpy.zeros.html) function.
1. The type of a numpy arra is ``ndarray``.

<!-- grade_id: cell-18fe0c61360d94af, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
plt.plot(tt, d1a(fp1(tt)))  
```

<!-- grade_id: cell-243ed50368c1b1c9, points: 3, grade: true, solution: false -->
```python
assert_equal(len(d1a(fp1(tt), h=0.1)), len(tt))
```

<!-- grade_id: cell-caa465425e04202b, points: 3, grade: true, solution: false -->
```python
assert_true(isinstance(d1a(fp1(tt), h=0.1), np.ndarray))
```

<!-- grade_id: cell-b31447c71891b333, points: 2, grade: true, solution: false -->
```python
assert_raises(TypeError, d1a, [1,2,3,4])
```

<!-- grade_id: cell-164d8329853bb8b0, points: 2, grade: true, solution: false -->
```python
assert_raises(TypeError, d1a, np.arange(-10, 10, 0.1), h=[0.1,0.1])
```

<!-- grade_id: cell-afc2e88a493f9788, points: 2, grade: true, solution: false -->
```python
assert_raises(ValueError, d1a, np.arange(-10, 10, 0.1), h=-0.1)
```

<!-- grade_id: cell-3ca6e7a29f7487ff, points: 13, grade: true, solution: false -->
```python

```

**Problem 2. (15 points)** Repeat problem 2 but create a function ``d1b`` that uses a centered difference ($f_i' \approx \frac{f_{i+1}-f_{i-1}}{2h}$). Again make sure that the resulting signal has the same length as the input signal and use slicing rather than a for loop to make the computation. Test the function with ``fp1``.

<!-- grade_id: cell-55e9e569ea973420, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
plt.plot(tt, d1b(fp1(tt), h=0.1))  
```

<!-- grade_id: cell-ac35996b8cd66fe3, points: 1, grade: true, solution: false -->
```python
assert_equal(len(d1b(fp1(tt), h=0.1)), len(tt))
```

<!-- grade_id: cell-7b09e3857ab89e6c, points: 1, grade: true, solution: false -->
```python
assert_true(isinstance(d1b(fp1(tt), h=0.1), np.ndarray))
```

<!-- grade_id: cell-ca4b1211b6642431, points: 1, grade: true, solution: false -->
```python
assert_raises(TypeError, d1b, [1,2,3,4])
```

<!-- grade_id: cell-e5f5b19536befbf7, points: 1, grade: true, solution: false -->
```python
assert_raises(TypeError, d1b, np.arange(-10, 10, 0.1), h=[0.1,0.1])
```

<!-- grade_id: cell-43a24511f5b6424c, points: 1, grade: true, solution: false -->
```python
assert_raises(ValueError, d1b, np.arange(-10, 10, 0.1), h=-0.1)
```

<!-- grade_id: cell-46f990e2eb4b9a83, points: 10, grade: true, solution: false -->
```python

```

**Problem 3. (30 points)**

Mathematicians are critical of the simple differences we have been using to approximate derivatives because they are inaccurate. More accurate approximations can be derived based on [Taylor series expansions](https://en.wikipedia.org/wiki/Taylor_series#Definition) of functions. A common approximation of the first derivative is:

$$
f_i' \approx \frac{-f_{i+2}+8f_{i+1}-8f_{i-1}+f_{i-2}}{12h}.
$$

This approximation to the derivative has an error an the order of the fourth power of the step size ($O(h^4)$).
(Strickwerda, *Finite Difference Schems and Partial Differential Equations,* p. 69)

Write a function ``d1c`` that implements the above approximation to the first derivative. Provide the same type checking (e.g. numpy array, number) as with problems 1 and 2. The function should have a docstring.

Test the function with fp1.


<!-- grade_id: cell-1b33724eb7917d85, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
plt.plot(tt, d1c(fp1(tt)))
```

<!-- grade_id: cell-f2f60265592a74e8, points: 1, grade: true, solution: false -->
```python
assert_equal(len(d1c(fp1(tt), h=0.1)), len(tt))
```

<!-- grade_id: cell-0ebe87833e860ad0, points: 1, grade: true, solution: false -->
```python
assert_true(isinstance(d1c(fp1(tt), h=0.1), np.ndarray))
```

<!-- grade_id: cell-5af6c529cea835dc, points: 1, grade: true, solution: false -->
```python
assert_raises(TypeError, d1c, [1,2,3,4])
```

<!-- grade_id: cell-84d812a44a9a9de4, points: 1, grade: true, solution: false -->
```python
assert_raises(TypeError, d1c, np.arange(-10, 10, 0.1), h=[0.1,0.1])
```

<!-- grade_id: cell-40e2602429ad99b4, points: 1, grade: true, solution: false -->
```python
assert_raises(ValueError, d1c, np.arange(-10, 10, 0.1), h=-0.1)
```

<!-- grade_id: cell-fb2e77f31a26abdd, points: 25, grade: true, solution: false -->
```python

```

```python
def relative_error(signal, true_signal, eps=0.1):
    """
    Computes the relative error between the estimated signal and true signal. Relative error is returns as
    100*(signal-true_signal)/true_signal.
    
    In order to avoid divide by zero issues, we assign a zero if the true_signal value is less than eps
    Arguments:
        signal: a 1D numpy array
        true_signal: a 1D numpy array
        eps: a floating point number indicating deviation from zero
    Returns:
        A numpy array with the relative error at each point
    
    """
    
    return np.where(np.abs(true_signal)>eps, 100*((signal-true_signal)/true_signal), 0)

```


To get a handle on the amount of error for each derivative, use the ``relative_error`` function to compute the error between the true derivative value (dfp1) and the derivative of the signal computed with each of the three functions (``d1a``, ``d1b``, ``d1c``).

Ignore the end points (where derivatives are not computed).

Create two scatter plots with matplotlib that shows the relative error of each derivative.

In the first scatter plot, plot the relative error of each derivative vs time (``tt``). In the seconed plot, plot the relative error of each derivative vs the true derivative. 

Set the ylimit for the plots so that the figure illustrates the relative performance of each technique.

```python
dc = d1c(fp1(tt), h=0.1)
db = d1b(fp1(tt), h=0.1)
da = d1a(fp1(tt), h=0.1)
dtrue = dfp1(tt)
```

```python
start = 4# 45
end = -4 #int(len(tt)*0.8)
f, ax = plt.subplots()
f.set_size_inches(6,6)
ax.scatter(fp1(tt)[start:end], relative_error(dc, dtrue)[start:end], marker='+', color='r')
ax.scatter(fp1(tt)[start:end], relative_error(db, dtrue)[start:end], marker='.', color='green')
ax.scatter(fp1(tt)[start:end], relative_error(da, dtrue)[start:end], marker='2', color='b')

plt.ylim(-10, 10)
```

```python
start = 4
end = -4 #int(len(tt)*0.8)
f, ax = plt.subplots()
f.set_size_inches(6,6)
ax.scatter(tt[start:end], relative_error(dc, dtrue)[start:end], marker='+', color='r')
ax.scatter(tt[start:end], relative_error(db, dtrue)[start:end], marker='.', color='green')
ax.scatter(tt[start:end], relative_error(da, dtrue)[start:end], marker='2', color='b')

#plt.plot(100*abs(dtrue[start:end]-db[start:end])/abs(dtrue[start:end]), 'g')
#plt.plot(100*abs(dtrue[start:end]-da[start:end])/abs(dtrue[start:end]), 'b')

plt.ylim(-10, 10)
```

**Problem 6 (40 points):**

A pharmokinetic model of drug absorption and clearance is given by

$$
y = \frac{Ab}{b-k}\left(e^{-kt}-e^{-bt}\right)
$$

where $k$ is a time constant that describes how the drug is cleared from the body (e.g. through urine) and $b$  is the clearance rate for the digestive system with $b > k$ typically.

For $A = 200$, $k = 0.35$, and $b = 0.46$ we want to find when the drug concentration is a maximum. Sympy provides [functions](http://docs.sympy.org/dev/modules/solvers/solvers.html#algebraic-equations) for finding the roots of an equation (when the equation equals zero) symbolically.

Compare the difference between the maximum dose time you find using sympy and what you obtain numerically using the function ``d1c`` defined above and finding the position in the resulting numpy array closest to zero. Use the numpy ``arange`` function with various step sizes (``h``). 

**Hints:**

1. Because ``d1c`` does not compute the derivatives near the endpoints of the array, you need to exclude these points when searching for the minimum.
1. Use the [numpy.argmin](http://docs.scipy.org/doc/numpy/reference/generated/numpy.argmin.html) function to find the minimum location in the derivative array and then use that index to find the value in the time array.
1. Because the derivative is both positive and negative, use the absolute value of the derivative to find the location nearest to zero.
1. Plot the derivative function and select reasonable limits to search for the minimum. (You need to be sure to exclude the non-computed points from the search.)
1. Remember ``d1c`` needs to be passed ``h``.

```python
from scipy.optimize import newton
from sympy.solvers import solve
```

<!-- grade_id: cell-b5b153a87a9ce327, points: 0, grade: false, solution: true -->
```python
ts = symbols('ts')
def kinetics(t, A=200, b=0.46, k=0.35):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-ec5d8d2942d78fce, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(kinetics(10), 16.8490035122456)
```

<!-- grade_id: cell-acb03fece91946a1, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(kinetics(5, A=195, b=0.32, k=0.2), 86.3111200519292)
```

```python
fkinetics = lambdify(ts, kinetics(ts), "numpy")
```

#### Plot the function to make sure everything looks correct

```python
plot(kinetics(ts), (ts,0, 20))
```

<!-- grade_id: cell-c581c86bd66fc373, points: 0, grade: false, solution: true -->
```python
def symbolic_roots(func, times):
    # YOUR CODE HERE
    raise NotImplementedError()
symbolic_roots(kinetics, ts)
```

<!-- grade_id: cell-05420e35f91efaa0, points: 10, grade: true, solution: false -->
```python

```

#### Create a time array

```python
h = 0.01
tn = np.arange(0, 20, h)
```

```python
plt.plot(tn, fkinetics(tn))
```

<!-- grade_id: cell-8ff4e7b89e907a30, points: 0, grade: false, solution: true -->
```python
def numeric_roots(func, ts, h=0.01, search_low=2, search_high=500):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c8c4ee3a57dc9e6c, points: 20, grade: true, solution: false -->
```python

```

```python

```